# Timeout continuation test

Tests, if the grading of all ngbrader-tests continue, even if there is a timeout during one single test.  
The test has to be auto-graded.  

One test in one file should fail.  
The overall grading should result in 40 points.

<!-- grade_id: cell-autograde-visible, points: 10, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;

System.out.println("Testing A");

Assertions.assertTrue(true);
```