# Timeout continuation test

Tests, if the grading of all ngbrader-tests continue, even if there is a timeout during one single test.  
The test has to be auto-graded.  

One test in one file should fail.  
The overall grading should result in 40 points.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
// BEGIN SOLUTION
// END SOLUTION
int timeoutSecondsPass = 1;
int timeoutSecondsFail = 400;
```

```java
boolean doSomething(int timeoutSeconds) throws java.lang.InterruptedException {
    long time = System.currentTimeMillis();
    long end = time + (timeoutSeconds * 1000);
    System.out.println("");
    while(System.currentTimeMillis() < end) {
        System.out.print(((end-time)/1000) + " ");
        Thread.sleep(5000);
        time = System.currentTimeMillis();
    }
    System.out.println("");
    return true;
}
```

<!-- grade_id: cell-autograde-visible-1, points: 1, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;

System.out.println("Testing C");

System.out.println("Waiting " + timeoutSecondsPass + " seconds");
Assertions.assertTrue(doSomething(timeoutSecondsPass));
System.out.println("Test 1 successfully completed");
```

<!-- grade_id: cell-autograde-visible-2, points: 1, grade: true, solution: false -->
```java
System.out.println("Waiting " + timeoutSecondsPass + " seconds");
Assertions.assertTrue(doSomething(timeoutSecondsPass));
System.out.println("Test 2 successfully completed");
```

<!-- grade_id: cell-autograde-visible-3, points: 1, grade: true, solution: false -->
```java
System.out.println("Waiting " + timeoutSecondsFail + " seconds");
Assertions.assertTrue(doSomething(timeoutSecondsFail));
System.out.println("Test 3 successfully completed");
```

<!-- grade_id: cell-autograde-visible-4, points: 1, grade: true, solution: false -->
```java
System.out.println("Waiting " + timeoutSecondsPass + " seconds");
Assertions.assertTrue(doSomething(timeoutSecondsPass));
System.out.println("Test 4 successfully completed");
```

<!-- grade_id: cell-autograde-visible-5, points: 1, grade: true, solution: false -->
```java
System.out.println("Waiting " + timeoutSecondsPass + " seconds");
Assertions.assertTrue(doSomething(timeoutSecondsPass));
System.out.println("Test 5 successfully completed");
```