# [DELETE THIS CELL] Template

- If cell types are not displayed, select in the menu: _View_ -> _Cell Toolbar_ -> _Create Assignment_.
- [Markdown](https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html)
- [JUnit Assertions](https://junit.org/junit5/docs/current/api/org/junit/jupiter/api/Assertions.html)
- [Maven Repository](https://mvnrepository.com/)
- [Apache Commons IO FileUtils](https://commons.apache.org/proper/commons-io/javadocs/api-release/index.html?org/apache/commons/io/FileUtils.html)

```java
%maven commons-io:commons-io:2.6
import org.apache.commons.io.FileUtils;
import org.apache.commons.io.filefilter.TrueFileFilter;


// [DELETE THIS CELL]


/* Distribution of files
 *
 * The directory for shared files is /srv/distribution/.
 * Students can read, but not write.
 */

// Read example
// Can be used in assignments.
File file = new File("/srv/distribution/hello-world.txt");
String string = FileUtils.readFileToString(file, "UTF-8");
System.out.println(string);

// Write example
// Only possible for nbgadmin.
// Code can be used to create files for assignments.
file = new File("/srv/distribution/tmp.txt");
FileUtils.writeStringToFile(file, "Test", "UTF-8");
string = FileUtils.readFileToString(file, "UTF-8");
file.delete();
System.out.println(string);

// List files check
Collection<File> files = FileUtils.listFiles(new File("/srv/distribution/"), TrueFileFilter.INSTANCE, TrueFileFilter.INSTANCE);
System.out.println();
for (File f : files) {
	System.out.println(f);
}
```

# Numeric ASCII representation

Implement a method wich converts a String to the numeric ASCII representation of its single characters.

#### Example

<table>
    <tr>
        <th></th>
        <th><p align="left">Input</p></th>
        <th><p align="left">Output</p></th>
    </tr>
    <tr>
        <td><p align="left"><b>Type</b></p></td>
        <td><p align="left">String</p></td>
        <td><p align="left">int array</p></td>
    </tr>
    <tr>
        <td><p align="left"><b>Value</b></p></td>
        <td><p align="left">Hi</p></td>
        <td><p align="left">72, 105</p></td>
    </tr>
</table>

#### Hints

- The [String class](https://docs.oracle.com/en/java/javase/11/docs/api/java.base/java/lang/String.html) provides methods for convertion.

#### Notes

- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
int[] toNumericAscii(String string) {
    int[] result = new int[string.length()];

    // BEGIN SOLUTION
    
    /* Notes for lecturers:
     *
     * If possible, the return variable should be initialized outside of the solution.
     * It avoids unnecessary error messages, if the solution part is removed.
     */

    char[] chars = string.toCharArray();
    for (int i = 0; i < chars.length; i++) {
        result[i] = chars[i];
    }
    // END SOLUTION

    return result;
}
```

# Evaluation

- Run the following cell to test your implementation.
- You can ignore the cells afterwards.

<!-- grade_id: cell-autograde-visible, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;
try {

    Assertions.assertEquals(105, toNumericAscii("Hi")[1]);
    Assertions.assertArrayEquals( new int[]{ 72, 105 }, toNumericAscii("Hi") );
    Assertions.assertTrue(72 == toNumericAscii("Hi")[0]);
    
    System.out.println("Test(s) successfully completed");
} catch(AssertionFailedError e) { System.err.print(e); }
```

<!-- grade_id: cell-autograde-1, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

/* Notes for lecturers:
 *
 * For the cell above 'cell-autograde-visible', there should be no points.
 * Use this cell instead to create a similar test.
 * This avoids to just copy and return the exact test solution.
 *
 * Overview of Assertions:
 * https://junit.org/junit5/docs/current/api/org/junit/jupiter/api/Assertions.html
 */

Assertions.assertEquals(111, toNumericAscii("Ho")[1]);
Assertions.assertArrayEquals( new int[]{ 72, 111 }, toNumericAscii("Ho") );
Assertions.assertTrue(72 == toNumericAscii("Ho")[0]);

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-2, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-3, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-4, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-5, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```