# Exercise: Numeric ASCII representation

(24 points)

Implement a method which converts a String to the numeric ASCII representation of its single characters.

#### Example

<table>
    <tr>
        <th></th>
        <th><p align="left">Input</p></th>
        <th><p align="left">Output</p></th>
    </tr>
    <tr>
        <td><p align="left"><b>Type</b></p></td>
        <td><p align="left">String</p></td>
        <td><p align="left">int array</p></td>
    </tr>
    <tr>
        <td><p align="left"><b>Value</b></p></td>
        <td><p align="left">Hi</p></td>
        <td><p align="left">72, 105</p></td>
    </tr>
</table>

#### Hints

- The [String class](https://docs.oracle.com/en/java/javase/11/docs/api/java.base/java/lang/String.html#toCharArray()) provides methods for conversion.
- Please read the [notes](#Notes) at the end of this notebook.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
public class Exercise {

    /**
     * Converts a String to the numeric ASCII representation of its single
     * characters.
     */
    public int[] toNumericAscii(String string) {
        int[] result = new int[string.length()];

        // BEGIN SOLUTION

        /*
         * Notes for lecturers:
         *
         * Markdown examples for the cell above:
         * https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html
         * 
         * Insert your solution inside one or several marked solution blocks. These
         * blocks will be removed for students.
         *
         * If possible, a return variable should be defined before this marked solution
         * block. It avoids unnecessary error messages, if the solution block is removed
         * for the students. A null value also works.
         * 
         * You can remove this comment.
         */

        char[] chars = string.toCharArray();
        for (int i = 0; i < chars.length; i++) {
            result[i] = chars[i];
        }

        // END SOLUTION

        return result;
    }
}
```

# Evaluation

- Run the following cell to test your implementation.
- You can ignore the empty cells afterwards.

<!-- grade_id: cell-autograde-visible, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

void test(String inputString, int[] expectedResult, String testDescription) {
    try {
        long time = System.currentTimeMillis();

        Assertions.assertArrayEquals(expectedResult, new Exercise().toNumericAscii(inputString),
                                     testDescription);

        time = System.currentTimeMillis() - time;
        System.out.println("Test(s) successfully completed. Calculation took " + time + " ms.");
    } catch (AssertionFailedError e) {
        System.err.println("Your solution returned an unexpected result. ");
        throw e;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error. ");
        throw e;
    }
}

test("Hi", new int[] { 72, 105 }, "Testing 'Hi'");
```

<!-- grade_id: cell-autograde-1, points: 14, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

/* Notes for lecturers:
 *
 * For the cell above 'cell-autograde-visible', there should be no points.
 * Use this cell instead to create a test case similar to the visible test.
 * This avoids students to just return the exact solution of the visible test.
 *
 * Use the last parameter of assertions to provide meaningful test descriptions.
 * A description "Test no. X" also helps in giving feedback.
 *
 * Feel free to change the parameters of the method test(...).
 * 
 * Overview of Assertions:
 * https://junit.org/junit5/docs/current/api/org/junit/jupiter/api/Assertions.html
 */

test("Ho", new int[] { 72, 111 }, "Testing 'Ho'");

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-2, points: 10, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

test("HiHo", new int[] { 72, 105, 72, 111 }, "Testing 'HiHo'");

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-3, points: 0, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-4, points: 0, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-5, points: 0, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

// END HIDDEN TESTS
```

# Notes

- You are free to use a different IDE to develop your solution. However, you have to copy the solution into this notebook to submit it.
- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ > _View_ > _Toggle Line Numbers_ for debugging.
  - Check _Menu_ > _Help_ > _Keyboard Shortcuts_.
- Known issues
  - All global variables will be set to void after an import.
  - Missing spaces around `%` (modulo operator) can cause unexpected errors so please make sure that you have added spaces around every `%` character.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Make sure that all necessary imports are listed at the beginning of your cell.
  - Run a final check of your solution by
    - Click the _fast forward arrow_ in the tool bar (or _Menu_ > _Kernel_ > _Restart & Run All_). Wait fo the kernel to restart and execute all cells (all executable cells should have numbers in front of them instead of a `[*]`).
    - Check all executed cells for errors. If an exception is thrown, please check your code. Note that although the error might look cryptic, until now we never encounter that an exception was caused without a valid reason inside of the submitted code. A good way to check the code is to copy the solution into a new class in your favorite IDE and check
      - errors reported by the IDE
      - imports the IDE adds to your code which might be missing in your submission, here
  - Finally, choose _Menu_ > _File_ > _Close and Halt_.
  - Do not forget to submit your solution in the _Assignments_ view.