## Multiple choice

# Which of the possible answers are colors?

Set the variables to **true**, if they represent a **correct** answer.

#### Notes

- **Do not change the names of the variables**.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
boolean cat   = false;
boolean green = false;
boolean red   = false;
boolean dog   = false;

// BEGIN SOLUTION
green = true;
red = true;
// END SOLUTION
```

# Evaluation

- You can ignore the following cells.

<!-- grade_id: cell-autograde-visible, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;
```

<!-- grade_id: cell-autograde-all, points: 3, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

/* Notes for lecturers:
 *
 * If you want to award points ONLY if ALL solutions are correct:
 * - Set all points to this cell.
 * - Set the points in other autograder-test-cells to 0.
 * - Do NOT delete other autograder-test-cells.
 *   They will provide feedback to the students.
 *
 * If you want to award points ONLY for EACH single solution:
 * - Set the points in this cell to 0.
 * - Set the points in other autograder-test-cells to 1 (or greater, if a greater sum is needed).
 * - You can remove the assertions in this cell.
 *
 * If you want to have a combination:
 * - Set points for all autograder-test-cells.
 *
 * Example configuration:
 * - 3+1+1+1+1 = 7 points for a perfect solution.
 * - 0+1+1+1+0 = 3 points, if 'dog is a color' is the opinion of the student.
 *
 * Overview of Assertions:
 * https://junit.org/junit5/docs/current/api/org/junit/jupiter/api/Assertions.html
 */

Assertions.assertFalse(cat, "general test");
Assertions.assertTrue(green, "general test");
Assertions.assertTrue(red, "general test");
Assertions.assertFalse(dog, "general test");

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-1, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

Assertions.assertFalse(cat, "cat");

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-2, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

Assertions.assertTrue(green, "green");

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-3, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

Assertions.assertTrue(red, "red");

// END HIDDEN TESTS
```

<!-- grade_id: cell-autograde-4, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
// BEGIN HIDDEN TESTS

Assertions.assertFalse(dog, "dog");

// END HIDDEN TESTS
```