# Time test

This tests the maximum execution time.  
It is defined by nbgrader option `ExecutePreprocessor.timeout`.  
The test has to be executed as an assignment.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
// BEGIN SOLUTION
// END SOLUTION
int timeoutSeconds = 300;
boolean testShouldPass = false;
```

```java
boolean doSomething(int timeoutSeconds) {
    long time = System.currentTimeMillis();
    long end = time + (timeoutSeconds * 1000);
    while(System.currentTimeMillis() < end) {
        time = System.currentTimeMillis();
    }
    return true;
}
```

<!-- grade_id: cell-autograde-visible, points: 1, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;

if(testShouldPass) {
    System.out.println("Waiting " + (timeoutSeconds - 20) + " seconds");
    Assertions.assertTrue(doSomething(timeoutSeconds - 20));
} else {
    System.out.println("Waiting " + (timeoutSeconds + 10) + " seconds");
    Assertions.assertTrue(doSomething(timeoutSeconds + 10));
}
System.out.println("Test(s) successfully completed");
```