# Example assignment for Java and nbgrader

You can use Markdown to provide information to the students.

* Jupyter notebook command help: [Ctrl] + [Shift] + [p]
* [Nbgrader documentation](https://nbgrader.readthedocs.io/)
* [Java kernel documentation](https://github.com/SpencerPark/IJava)
* [Java kernel examples](https://github.com/SpencerPark/ijava-binder)
* [DICE nbgraderutils](https://github.com/dice-group/nbgraderutils)

### Cell types in nbgrader

Lecturers have to define the type of a cell, e.g. Markdown or Code. You can choose this at the top menue.

Additionally, you can set the points for an answer and create manually or autograded answers. To do so, click on "View" -> "Cell Toolbar" -> "Create Assignment".

Have a look at the examples below.

## 1. Manually graded text task
Who is the best lecturer in the world?

<!-- grade_id: cell-ef07e58f8575b7d2, points: 1, grade: true, solution: true -->


## 2. Manually graded programming task

Write a method call which prints a greetings string.

<!-- grade_id: cell-8bd565595479558a, points: 2, grade: true, solution: true -->
```java

```

## 3. Autograded programming task

Write a method which calculates the product of two integers.

<!-- grade_id: cell-b65aa29f897ebc4d, points: 0, grade: false, solution: true -->
```java
public int product(int a, int b) {
    // BEGIN SOLUTION
    // The correct solution of a lecturer.
    // This part will be removed for the students.
    // You can use this solution to create tests in the next cell.
    return a * b;
    // END SOLUTION
    
    // This would be an example for an incorrect student answer failing at the public test.
    // return 2;
    
    // This would be an example for an incorrect student answer, passing the public test.
    // return ( (a + 1) * (b - 1) );
}
```

<!-- grade_id: cell-41099f4980cf7022, points: 4, grade: true, solution: false -->
```java
public void assertTrue(boolean assertion){if(!assertion){throw new RuntimeException("Test failed");}}

// Execute these tests to check your solution.
assertTrue(product(1,2) == 2);
assertTrue(product(5,6) == 30);



// BEGIN HIDDEN TESTS
public boolean assertTrue(boolean assertion, String comment) {
    System.out.println(assertion ? "Passed: " + comment : "Failed: " + comment);
    return assertion;
}
boolean allTrue = true;

// Hidden tests of lecturers
// This part will be removed for the students.
// For every assertion, it is printed, if the related test passes or fails.
// Additionally, if one of the tests fails, an exception is thrown for autograding.
allTrue = allTrue & assertTrue( product(2,3) == 6, "Hidden test 1");
allTrue = allTrue & assertTrue( product(3,3) == 9, "Hidden test 2 (square)");
allTrue = allTrue & assertTrue( product(4,5) == 20, "Hidden test 3");

if(!allTrue){throw new RuntimeException("Hidden test failed");}
// END HIDDEN TESTS
```

Output:
```
Passed: Hidden test 1
Passed: Hidden test 2 (square)
Passed: Hidden test 3
```

## 4. Autograded programming task with library

* Task 1: Complete the methods levenshtein() and jaccard() to compute the respective distances.
* Task 2: Compute the levenshtein distance and the jaccard distance for the words A and B.
* Hint: [Documentation of the package Apache Commons Text](https://commons.apache.org/proper/commons-text/)

```java
%%loadFromPOM
<dependency>
   <groupId>org.apache.commons</groupId>
   <artifactId>commons-text</artifactId>
   <version>1.4</version>
</dependency>
```

```java
// Import the needed packacke members
import org.apache.commons.text.similarity.*;

// The words to compare
String wordA = "house";
String wordB = "mouse";

// Variables to be set
int levenshteinDistance = 0;
double jaccardDistance = 0;
```

<!-- grade_id: cell-2c5cfbdf011615ee, points: 0, grade: false, solution: true -->
```java
public int levenshtein(String wordA, String wordB) {
    // BEGIN SOLUTION
    return LevenshteinDistance.getDefaultInstance().apply(wordA, wordB);
    // END SOLUTION
}

public double jaccard(String wordA, String wordB) {
    // BEGIN SOLUTION
    return new JaccardDistance().apply(wordA, wordB);
    // END SOLUTION
}

// BEGIN SOLUTION
levenshteinDistance = levenshtein(wordA, wordB);
jaccardDistance = jaccard(wordA, wordB);
// END SOLUTION
```

<!-- grade_id: cell-92c583222889bae3, points: 8, grade: true, solution: false -->
```java
// Execute these tests to check your solution.
assertTrue(levenshtein("dog", "mouse") == 4);
assertTrue(jaccardDistance != 0);

// BEGIN HIDDEN TESTS

// Hidden tests of lecturers
// This part will be removed for the students.
// For every assertion, it is printed, if the related test passes or fails.
// Additionally, if one of the tests fails, an exception is thrown for autograding.
boolean allTrue = true;
allTrue = allTrue & assertTrue( levenshteinDistance == 1, "levenshtein A B" );
allTrue = allTrue & assertTrue( jaccardDistance > 0.31, "jaccard A B >" );;
allTrue = allTrue & assertTrue( jaccardDistance < 0.35, "jaccard A B <" );
allTrue = allTrue & assertTrue( levenshtein("computer","cpu") == 5, "levenshtein method" );
allTrue = allTrue & assertTrue( jaccard("computer","cpu") > 0.60, "jaccard method >" );
allTrue = allTrue & assertTrue( jaccard("computer","cpu") < 0.64, "jaccard method <" );

if(!allTrue){throw new RuntimeException("Hidden test failed");}
// END HIDDEN TESTS
```

Output:
```
Passed: levenshtein A B
Passed: jaccard A B >
Passed: jaccard A B <
Passed: levenshtein method
Passed: jaccard method >
Passed: jaccard method <
```