# Simple example for C++ and nbgrader

You can use Markdown to provide information to the students.

* Command help: [Ctrl] + [Shift] + [p]
* [Nbgrader documentation](https://nbgrader.readthedocs.io/)
* [C++ kernel documentation](https://xeus-cling.readthedocs.io/)
* [Some DICE information](https://github.com/dice-group/nbgraderutils)

### Cell types in nbgrader

Lecturers have to define the type of a cell, e.g. Markdown or Code. You can choose this at the top menue.

Additionally, you can set the points for an answer and create manually or autograded answers. To do so, click on "View" -> "Cell Toolbar" -> "Create Assignment".

Have a look at the examples below.

## 1. Manually graded text task
Who is the best lecturer in the world?

<!-- grade_id: cell-ef07e58f8575b7d2, points: 1, grade: true, solution: true -->
 
 

## 2. Manually graded programming task

Write a method call which prints a greetings string.  
Use the namespace "xcpp", the scope resolution operator "::", and the function "display".

Additionally, you can print strings using the namespace "std", the standard output stream "cout", and the insertion operator "<<".

```C++17
#include <xcpp/xdisplay.hpp>
#include <iostream>
```

<!-- grade_id: cell-8bd565595479558a, points: 2, grade: true, solution: true -->
```C++17

```

## 3. Autograded programming task

Write a method which calculates the product of two integers.

<!-- grade_id: cell-b65aa29f897ebc4d, points: 0, grade: false, solution: true -->
```C++17
int product(int a, int b)
{
    // BEGIN SOLUTION
    return a * b;
    // END SOLUTION
    
    // This would be an example for an incorrect student answer failing at the public test.
    // return 2;
    
    // This would be an example for an incorrect student answer, passing the public test.
    //return ( (a + 1) * (b - 1) );
}
```

```C++17
// This function is used in the next cell to test assertions
void assertTrue(bool assertion) {
    if(!assertion){
        throw std::runtime_error("Test failed");
    }
}
```

<!-- grade_id: cell-3318ea41805fb90a, points: 4, grade: true, solution: false -->
```C++17
// Execute these tests to check your solution.
assertTrue(product(1,2) == 2);
assertTrue(product(5,6) == 30);
```

<!-- grade_id: cell-41099f4980cf7022, points: 0, grade: true, solution: false -->
```C++17
// BEGIN HIDDEN TESTS
// This function is used in the next cell to test assertions
bool assertTrue(bool assertion, std::string comment)
{
    std::cout << (assertion ? "Passed: " + comment : "Failed: " + comment ) << '\n';
    return assertion;
}
// END HIDDEN TESTS
```

<!-- grade_id: cell-d27f94236d836581, points: 8, grade: true, solution: false -->
```C++17
// BEGIN HIDDEN TESTS
bool allTrue = true;

// Hidden tests of lecturers
// This part will be removed for the students.
// For every assertion, it is printed, if the related test passes or fails.
// Additionally, if one of the tests fails, an exception is thrown for autograding.
allTrue = allTrue & assertTrue( product(2,3) == 6, "Hidden test 1");
allTrue = allTrue & assertTrue( product(3,3) == 9, "Hidden test 2 (square)");
allTrue = allTrue & assertTrue( product(4,5) == 20, "Hidden test 3");

if(!allTrue){throw std::runtime_error("Hidden test failed");}
// END HIDDEN TESTS
```

Output:
```
Passed: Hidden test 1
Failed: Hidden test 2 (square)
Passed: Hidden test 3
Standard Exception: Hidden test failed
```