# Advanced example for Java and nbgrader

* Task 1: Complete the methods levenshtein() and jaccard() to compute the respective distances.
* Task 2: Compute the levenshtein distance and the jaccard distance for the words A and B.
* Hint: [Documentation of the package Apache Commons Text](https://commons.apache.org/proper/commons-text/)

```java
%%loadFromPOM
<dependency>
   <groupId>org.apache.commons</groupId>
   <artifactId>commons-text</artifactId>
   <version>1.4</version>
</dependency>
```

```java
// Import the needed packacke members
import org.apache.commons.text.similarity.*;

// The words to compare
String wordA = "house";
String wordB = "mouse";

// Variables to be set
int levenshteinDistance = 0;
double jaccardDistance = 0;
```

<!-- grade_id: cell-3b2ccd042ea41283, points: 0, grade: false, solution: true -->
```java
public int levenshtein(String wordA, String wordB) {
    // BEGIN SOLUTION
    return LevenshteinDistance.getDefaultInstance().apply(wordA, wordB);
    // END SOLUTION
}

public double jaccard(String wordA, String wordB) {
    // BEGIN SOLUTION
    return new JaccardDistance().apply(wordA, wordB);
    // END SOLUTION
}

// BEGIN SOLUTION
levenshteinDistance = levenshtein(wordA, wordB);
jaccardDistance = jaccard(wordA, wordB);
// END SOLUTION
```

<!-- grade_id: cell-55bfad1814baadea, points: 4, grade: true, solution: false -->
```java
public void assertTrue(boolean assertion){if(!assertion){throw new RuntimeException("Test failed");}}

// Execute these tests to check your solution.
assertTrue(levenshtein("dog", "mouse") == 4);
assertTrue(jaccardDistance != 0);



// BEGIN HIDDEN TESTS
public boolean assertTrue(boolean assertion, String comment) {
    System.out.println(assertion ? "Passed: " + comment : "Failed: " + comment);
    return assertion;
}
boolean allTrue = true;

// Hidden tests of lecturers
// This part will be removed for the students.
// For every assertion, it is printed, if the related test passes or fails.
// Additionally, if one of the tests fails, an exception is thrown for autograding.
allTrue = allTrue & assertTrue( levenshteinDistance == 1, "levenshtein A B" );
allTrue = allTrue & assertTrue( jaccardDistance > 0.31, "jaccard A B >" );;
allTrue = allTrue & assertTrue( jaccardDistance < 0.35, "jaccard A B <" );
allTrue = allTrue & assertTrue( levenshtein("computer","cpu") == 5, "levenshtein method" );
allTrue = allTrue & assertTrue( jaccard("computer","cpu") > 0.60, "jaccard method >" );
allTrue = allTrue & assertTrue( jaccard("computer","cpu") < 0.64, "jaccard method <" );

if(!allTrue){throw new RuntimeException("Hidden test failed");}
// END HIDDEN TESTS
```

Output:
```
Passed: levenshtein A B
Passed: jaccard A B >
Passed: jaccard A B <
Passed: levenshtein method
Passed: jaccard method >
Passed: jaccard method <
```