# Simple example for Java and nbgrader

You can use Markdown to provide information to the students.

* [Nbgrader documentation](https://nbgrader.readthedocs.io/)
* [Java kernel examples](https://github.com/SpencerPark/IJava)
* [Some DICE information](https://github.com/dice-group/nbgraderutils)

### Cell types in nbgrader

Lecturers have to define the type of a cell, e.g. Markdown or Code. You can choose this at the top menue.

Additionally, you can set the points for an answer and create manually or autograded answers. To do so, click on "View" -> "Cell Toolbar" -> "Create Assignment".

Have a look at the examples below.

## 1. Manually graded text task
Who is the best lecturer in the world?

<!-- grade_id: cell-ef07e58f8575b7d2, points: 1, grade: true, solution: true -->
The best lecturer in the world is: 

## 2. Manually graded programming task

Write a method call which prints a greetings string.

<!-- grade_id: cell-8bd565595479558a, points: 2, grade: true, solution: true -->
```java
// System...
```

## 3. Autograded programming task

Write a method which calculates the product of two integers.

<!-- grade_id: cell-b65aa29f897ebc4d, points: 0, grade: false, solution: true -->
```java
public int product(int a, int b) {
    // BEGIN SOLUTION
    // The correct solution of a lecturer.
    // This part will be removed for the students.
    // You can use this solution to create tests in the next cell.
    return a * b;
    // END SOLUTION
    
    // This would be an example for an incorrect student answer failing at the public test.
    // return 2;
    
    // This would be an example for an incorrect student answer, passing the public test.
    // return ( (a + 1) * (b - 1) );
}
```

<!-- grade_id: cell-41099f4980cf7022, points: 4, grade: true, solution: false -->
```java
public void assertTrue(boolean assertion){if(!assertion){throw new RuntimeException("Test failed");}}

// Execute these tests to check your solution.
assertTrue(product(1,2) == 2);
assertTrue(product(5,6) == 30);



// BEGIN HIDDEN TESTS
public boolean assertTrue(boolean assertion, String comment) {
    System.out.println(assertion ? "Passed: " + comment : "Failed: " + comment);
    return assertion;
}
boolean allTrue = true;

// Hidden tests of lecturers
// This part will be removed for the students.
// For every assertion, it is printed, if the related test passes or fails.
// Additionally, if one of the tests fails, an exception is thrown for autograding.
allTrue = allTrue & assertTrue( product(2,3) == 6, "Hidden test 1");
allTrue = allTrue & assertTrue( product(3,3) == 9, "Hidden test 2 (square)");
allTrue = allTrue & assertTrue( product(4,5) == 20, "Hidden test 3");

if(!allTrue){throw new RuntimeException("Hidden test failed");}
// END HIDDEN TESTS
```

Output:
```
Passed: Hidden test 1
Passed: Hidden test 2 (square)
Passed: Hidden test 3
```